# Lab 5:  Modeling (Misspecification, Empirical Bayes, and Prior Predictive Checking)

Welcome to the fifth Data 102 lab! 

This lab has three goals. You will:

1. Study the sensitivity of Bayesian estimators to errors in the chosen prior model. 

1. Fit the parameters of a prior by maximizing the evidence. 

1. Validate your prior by running a prior predictive check.

You will be asked to complete the code in this lab from scratch. You are welcome to use AI tools to help you code, but are responsible for the correctness of the code, and your own comprehension of its mechanics. You may be asked to replicate code snippets requested here, as pseudo-code, on an exam. Where asked to describe a plot result, or reason with a conclusion, double check any code outputs. You may be asked to provide similar description or reasoning for sample figures/outputs on an exam. 

# Question 1: Sensitivity to Prior Misspecification

This question is closely related to discussion 5. Please review that question before starting here. Note, since this problem reuses work from discussion 5, the questions on discussion 5 will be in scope for your second exam. 

In most Bayesian contexts we do not know our priors exactly. So, it is important to study how robust our inferential procedures are to errors in the prior (that is, to prior misspecification). 

Consider the setting you explored in discussion 5. You can observe a string of i.i.d. random variables $\{X_i\}_{i=1}^n$ drawn normally with unknown mean $\Theta$ and known variance $\sigma_x^2$. To estimate the mean you adopt the Bayesian model:

$$\Theta \sim \text{Normal}(\mu_{\theta},\sigma_{\theta}^2), \quad \quad \{X_i\}_{i=1}^n |\Theta = \theta \underset{\text{i.i.d}}{\sim} \text{Normal}(\theta,\sigma_x^2)$$

Then, under your model:

$$\Theta|\{X_i = x_i\}_{i=1}^n \sim \text{Normal}\left( a, b \right) $$

where:

$$\begin{aligned}
    & a = (1 - \lambda) \hat{\theta}_{\text{MLE}}(\{x_i\}_{i=1}^n) + \lambda \mu_{\theta}, \\
    & \hat{\theta}_{\text{MLE}}(\{x_i\}_{i=1}^n) = \frac{1}{n} \sum_{i=1}^n x_i, \\
    & \lambda = \frac{1/\sigma_{\theta}^2}{n/\sigma_x^2 + 1/\sigma_{\theta}^2} \\
    & b = \frac{1}{n/\sigma_x^2 + 1/\sigma_{\theta}^2}
\end{aligned} $$

You adopt the posterior mean as your estimator:

$$\hat{\theta}_{\text{mean}}(\{x_i\}_{i=1}^n) = \mathbb{E}[\Theta|\{X_i = x_i\}_{i=1}^n] =  (1 - \lambda) \hat{\theta}_{\text{MLE}}(\{x_i\}_{i=1}^n) + \lambda \mu_{\theta}$$

Suppose that, in truth:

$$\Theta \sim \text{Normal}(0,1). $$

Then your model is misspecified whenever $\mu_{\theta} \neq 0$ or $\sigma_{\theta} \neq 1$. Let's work out how much error we can admit before our posterior mean estimator is less reliable (has a larger risk) than the MLE estimator. 

## 1a) Checking the MLE

Confirm that, the MLE is the sample average, $\bar{x}_n$:
    
$$\hat{\theta}_{\text{MLE}}(\{x_i\}_{i=1}^n) = \bar{x}_n = \frac{1}{n} \sum_{i=1}^n x_i. $$

_Type your answer here, replacing this text._

## 1b) Generality

It may seem odd that we are starting from a fixed true model. Show that, if, in truth, $\Psi \sim \text{Normal}(m,v) $, we assumed $\Psi \sim \text{Normal}(\mu_{\psi},\sigma_{\psi}^2)$, and observed $\{Y_i\}_{i=1}^n|\Psi = \psi \sim \text{Normal}(\psi,\sigma_y^2)$, then we could recover the case introduced above via the change of coordinates $x = (y - m)/\sqrt{v}$. Apply this change of coordinates to express $\Theta$ as a function of $\Psi$. Show that the true prior for $\Theta$ is $\text{Normal}(0,1)$, and our assumed prior is $\text{Normal}((\mu_{\psi} - m)/\sqrt{v},\sigma_{\psi}^2/v)$. 

_Type your answer here, replacing this text._

So, we can assume, without loss of generality, that the true prior is $\Theta \sim \text{Normal}(0,1)$ and interpret:

1. $\mu_{\theta}$ as the standardized discrepancy in means, $ (\mu_{\psi} - m)/\sqrt{v}$, and
1. $\sigma_{\theta}$ as the ratio of the assumed and true standard deviations $\sigma_{\psi}/\sqrt{v}$. 

## 1c) Bias and Variance in the Estimators

Repeat steps (c) - (f) from discussion 5 to compute the Bayes risk of the posterior mean estimator and the MLE, averaging over the *true* prior model:

$$R_B(\hat{\theta}) = \mathbb{E}_{\Theta,\{X_i\}_{i=1}^n}[(\hat{\theta}(\{X_i\}_{i=1}^n) - \Theta)^2].  $$

Note, most of this analysis will be the same. However, the posterior mean estimator will now include an additional bias associated with the misspecification error.

### 1c (i) Bias

Find the Frequentist and Bayes bias in the MLE and posterior mean estimator, averaging over the *true* prior model, but drawing estimates using the assumed prior.

_Type your answer here, replacing this text._

### 1c (ii) Variance

Find the Frequentist and Bayes variance in the MLE and the posterior mean estimator, averaging over the *true* prior model, but drawing estimates using the assumed prior.

_Type your answer here, replacing this text._

### 1c (iii) Risk

Find the Bayes risk for each estimator, averaging over the *true* prior model, but drawing estimates using the assumed prior.

_Type your answer here, replacing this text._

## 1d) Risk Comparison

Show that, the difference between the Bayes risks for the MLE and the posterior mean estimator is:

$$R_B(\hat{\theta}_{\text{MLE}}) - R_B(\hat{\theta}_{\text{mean}}) = \left(1 - (1 - \lambda)^2 \right) R_B(\hat{\theta}_{\text{MLE}}) - \lambda^2 (\mu_{\theta}^2 + 1), \text{ where } R_B(\hat{\theta}_{\text{MLE}}) = \frac{\sigma_x^2}{n}$$



_Type your answer here, replacing this text._

Call this difference in risks $\Delta R$. The posterior mean estimator is more reliable than the MLE as long as $\Delta R > 0$. It is much more reliable if $\Delta R \gg 0$. It is less reliable if $\Delta R < 0$. 

In the space below, write a function that can compute $\Delta R$ given the parameters:

1. The standardized error in the estimated prior mean: $\mu_{\theta} = (\mu_{\psi} - m)/\sqrt{v}$

1. The ratio of the assumed prior standard deviation and true standard deviation: $\sigma_{\theta} = \sigma_{\psi}/\sqrt{v}$

1. The standard deviation in the sampled data: $\sigma_x$

1. The sample size $n$.

Your code should be vectorized so it can compute $\Delta R$ for many different input parameter combinations efficiently.

```python
# Prepare your code here
```

In the space below, generate a three plots, each showing $\Delta R$ as a function of $\mu_{\theta} \in [-5,5]$ and $\sigma_{\theta} \in (0.1,10]$ for $\sigma_x = 1$ and $n = 1$, $n = 10$, and $n = 100$. Use the following display settings:

1. Set your $x$-axis to $\mu_{\theta}$ and your $y$ axis to $\sigma_{\theta}$. Use a log 10 scale for $\sigma_{\theta}$.

1. Display the surface $\Delta R(\mu,\sigma)$ as a heatmap with overlaid contours. Add a colorbar for each plot.

1. Overlay, in bold, the contour corresponding to $\Delta R(\mu,\sigma) = 0$. 

```python
# Prepare your code here
```

## 1e) Conclusions

Whenever you generate a plot, you should reason with the result. In the space below, describe the plots you generated above, highlighting observations you think would be important to a user who wants to adopt the posterior mean estimator, but is concerned that they will use the wrong prior parameters. Make sure that your description answers the question: "*How large can the errors in the prior parameters be before the MLE is more reliable than the posterior mean estimator?*"

_Type your answer here, replacing this text._

Suppose that you select the correct prior variance, $\sigma_{\theta}^2 = 1$. Find the range of prior means, $\mu_{\theta}$ for which the posterior mean estimator is more reliable than the MLE.

_Type your answer here, replacing this text._

Suppose that you select the correct prior mean, $\mu_{\theta} = 0$. Find the range of prior variances, $\sigma_{\theta}^2$ for which the posterior mean estimator is more reliable than the MLE.

_Type your answer here, replacing this text._

You should notice that these answers are different. Given the correct variance, there is an upper and lower bound on the assumed means outside of which the MLE is more reliable. In contrast, given the correct prior mean, the posterior mean estimator is more reliable whenever $\sigma_{\theta}^2$ is greater than a lower bound. In other words, there is no upper bound on the error in the assumed variance past which we should use the MLE. We are safe to use the posterior mean estimator *no matter* how much we overestimate the prior variance. 

Justify this observation conceptually.

_Type your answer here, replacing this text._

# Question 2: Fitting Prior Parameters

Consider the following conjugate Bayesian model:

1. $\{\Lambda_i\}_{i=1}^n$ are drawn i.i.d. from a $\text{Gamma}(a,b)$ distribution:

    $$ \Lambda_i \geq 0, \quad f_{\Lambda_i}(\lambda) = \frac{b^a}{\Gamma(a)} \lambda^{a - 1} e^{-b \lambda} $$

1. $\{X_i|\Lambda_i = \lambda \}_{i=1}^n$ are drawn independently from a $\text{Poisson}(\lambda)$ distribution

    $$X_i \in \{0,1,2,...\}, \quad \text{Pr}(X_i = x|\Lambda_i = \lambda) = \frac{\lambda^x}{x!} e^{-\lambda}$$

You can observe the counts $\{X_i\}_{i=1}^n$, but not $\{\Lambda_i\}_{i=1}^n$. You would like to use your observations to estimate the prior parameters $a, b$ so that, given a new observation $X_{n+1} = x_{n+1}$, you could form posterior estimates to $\Lambda_{n+1}$. 

## 2a) Posterior Inference

Suppose, for now, that you had already fit for $a$ and $b$ (see below), and we observed $X_{n+1} = x_{n+1}$. Find the posterior distribution for $\Lambda_{n+1}|X_{n+1} = x_{n+1}$ and provide a formula for the maximum a posteriori estimator for $\Lambda_{n+1}$. 

_Type your answer here, replacing this text._

## 2b) The Evidence

In a Bayesian model the evidence is the marginal distribution of the data, marginalizing over the unknowns. 

Find the marginal distribution of $X_i$, marginalizing over $\Lambda_i$. Express your answer by completing the probability mass function:

$$\text{Pr}(X_i = x_i) = \underline{~~~~~~~~~~~~~~~~~~~~~~~~~~~~~} $$

To close your integral you may use the value of the gamma integral suggested by the normalization factor of the gamma distribution:

$$\int_{\lambda = 0}^{\infty} \lambda^{c - 1} e^{-d \lambda} d\lambda = \frac{\Gamma(c)}{d^c}. $$


_Type your answer here, replacing this text._

This distribution is also called the *prior predictive distribution* since it is the distribution we would use to predict the value of a new data point before observing any other data. Notice that its parameters are the parameters of the prior.

So, the prior predictive distribution, evaluated at an observed outcome, is, as a function of the prior parameters, the likelihood function for the prior parameters. It follows that, selecting the prior parameters that maximize the evidence is equivalent to selecting the maximum likelihood *prior* parameters. The prior parameters that maximize the evidence are the prior parameters that, if true, would make the observed data most likely!

The argumen above justifies the name "*evidence*". The larger the prior predictive distribution, evaluated at the observation, the more evidence we have in favor of our prior model (here, our choice of prior parameters. )

In the space below, complete the expressions:

$$\text{Pr}(X_1 = x_1,X_2 = x_2,...,X_n = x_n) = \underline{~~~~~~~~~~~~~~~~~~~~~~~~~~~~~}, \quad  \log(\text{Pr}(X_1 = x_1,X_2 = x_2,...,X_n = x_n)) = \underline{~~~~~~~~~~~~~~~~~~~~~~~~~~~~~}$$

_Type your answer here, replacing this text._

In the code cell below, write a function which can compute $\log(\text{Pr}(X_1 = x_1,X_2 = x_2,...,X_n = x_n))$ as a function of the observations $\{x_i\}_{i=1}^n$, and prior parameters $a$ and $b$.

```python
# Prepare your code here
```

## 2c) Empirical Bayes

Suppose that you observe the string of counts (sorted from smallest to largest):

$$\{X_i\}_{i=1}^{20} = \{0,3,6,6,8,12,13,13,14,15,16,20,20,21,24,24,27,27,30,40\} $$

Write a code below that finds the parameters $a$ and $b$ which maximize the likelihood of this observation. Note: it will be more stable to maximize the log evidence instead of the evidence itself. You may call a pre-built optimizer.

```python
# Prepare your code here
```

Return your best fit parameters below:

_Type your answer here, replacing this text._

## 2d) Model Validation (Prior Predictive Checking)

Next, we should check whether our best fit model could have plausibly generated the observed sample. 

In the space below, write a code that samples from the Bayesian model for arbitrary prior parameters $a$ and $b$. Your code should produce sample strings of length $n$ for any integer $n$.

```python
# Prepare your code here
```

Run a NHST to determine whether your best fit model could have plausibly generated the observed data. To run your NHST:

1. Adopt a one-sample [Kolmogorov-Smirnov test](https://en.wikipedia.org/wiki/Kolmogorov%E2%80%93Smirnov_test). The KS test applies since, under the joint model, the samples $\{X_i\}_{i=1}^n$ are drawn i.i.d. from the prior predictive distribution. Your KS statistic should equal the maximum absolute difference between the empirical cumulative distribution function of the observed samples and your best fit cumulative distribution function. 

1. Estimate the $p$-value for your observed KS statistic by sampling 10,000 samples of length 20 from your best fit model, then for each sampled data set:

    - refit for the best $a, b$ 
    
    - compute the associated KS statistic
    
    - then counting the fraction that are larger than or equal to the observed KS statistic.

In the space below, explain why we should refit $a ,b$ for each simulated dataset. If you are unsure, repeat the process with and without refitting and compare $p$-values. 

_Type your answer here, replacing this text._

Prepare your code below:

```python
# Prepare your code to compute the hypothetical CDF here
```

```python
# Prepare your code to compute the KS statistic between an empirical CDF and a hypothetical CDF here
```

```python
# Prepare your code for estimating the p-value by resampling here
```

_Report your estimated $p$-value and test result here by replacing this text._

If your estimated $p$-value was larger than 0.05, can you conclude that your model is correct? If not, what is the correct interpretation of your prior predictive check?

_Type your answer here, replacing this text._

To show that this validation check is capable of falsifying models either:

1. Vary the prior parameters until you find a pair where the test fails, or

1. Generate hypothetical data under a different model for $\Lambda_i$ that is far from Gamma (e.g. $\Lambda_i \sim 1 + 100 \times \text{Bern}(0.5)$). You may generate hypothetical data with $n > 20$ so that your test has higher power. Note: if you generate new data, you should refit for $a$ and $b$.

```python
# Prepare your code here
```

Comment on the power of this test. Given only twenty samples can we sharply resolve the prior, or, can you find a pair of highly distinct priors that both pass the test?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Well done!
You've reached the end of the lab! Make sure you double check your work and make sure that you've answered all the written portions of the lab.

Before you submit to Pensieve, make sure you pass all the autograded portions of this lab. **Run the last cell to generate a zip file of your lab submission.**

To submit your lab to Pensieve, submit the zip file to Lab 5 Code.

```python
import matplotlib.image as mpimg

img = mpimg.imread('baby_rock_hyrax.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(pdf=False, force_save=True)
```

Tests for q1a (points: 1)
```
>>> n_test_array = [10, 100, 1000]
>>> x_test_array = [10, 34, 852]
>>> res_array = [trivial_theta_estimate(n, x_test_array[i]) for i, n in enumerate(n_test_array)]
>>> assert np.allclose(res_array, [1.0, 0.34, 0.852])
```


Tests for q1d (points: 1)
```
>>> N_test = 100
>>> X_test = 20
>>> alpha_test_array = [1, 10, 100]
>>> beta_test_array = [1, 10, 100]
>>> inputs = list(itertools.product(alpha_test_array, beta_test_array))
>>> outputs = [posterior_mean_estimate(N_test, X_test, *inp) for inp in inputs]
>>> assert np.allclose(outputs, [0.20588235294117646, 0.1891891891891892, 0.1044776119402985, 0.2702702702702703, 0.25, 0.14285714285714285, 0.5970149253731343, 0.5714285714285714, 0.4])
```


Tests for q2a_i (points: 1)
```
>>> logger.setLevel(logging.ERROR)
>>> model, trace = approximate_inference_MCMC(10, 20)
>>> thetas = trace['theta']
>>> logger.setLevel(logging.INFO)
>>> assert thetas.shape == (2000, 12)
```


Tests for q2c (points: 1)
```
>>> logger.setLevel(logging.ERROR)
>>> nreps = 2
>>> answers = [[0.3], [0.26], [0.34], [0.31], [0.29], [0.33], [0.32, 0.33], [0.28], [0.23], [0.37], [0.43], [0.35]]
>>> checks = []
>>> for _ in range(nreps):
...     posterior_estimates_test = empirical_posterior_mean_estimates(10, 25)
...     checks.append([np.round(posterior_estimates_test[i], 2) in answers[i] for i in range(12)])
>>> logger.setLevel(logging.INFO)
>>> assert np.any(checks)
```


Tests for q2e_i (points: 1)
```
>>> nreps = 2
>>> logger.setLevel(logging.ERROR)
>>> answers = [0.38, 0.3, 0.44, 0.38, 0.36, 0.42, 0.42, 0.36, 0.28, 0.48, 0.56, 0.44]
>>> checks = []
>>> for _ in range(nreps):
...     model_test, trace_test = approximate_inference_asympotmatic_MCMC(5, 10)
...     post_samples_test = trace_test['theta']
...     estimates = np.mean(post_samples_test, axis=0)
...     rounded_estimates = np.round(estimates / 2, 2) * 2
...     checks.append(np.all([rounded_estimates[i] == answers[i] for i in range(len(rounded_estimates))]))
>>> logger.setLevel(logging.INFO)
>>> assert np.any(checks)
```
